# Aula 2 - Exercícios Extra - Sintaxe de Python e NumPy Array

```python
#@title Init Cell - Execute-me!!!
import sys
from datetime import datetime
import requests
import time
from requests.utils import quote
url = 'https://docs.google.com/forms/d/e/1FAIpQLSeASSC8-w8FmfodZ4lBnuSEAvYuE4vatIBowLIREG1f-2pIpA/formResponse?usp=pp_url&entry.1986154915=mbacd&entry.513694412=2021&entry.1914621244=CienciaDeDados'

def format_values(values):
    return {
        "entry.1269959472": values['student_id'],
        "entry.1799867692": str(values['exercise_number']).replace(".", "_"),
        "entry.886231469": values['exercise_score'],
        "entry.1491599254": values['exercise_extra'],
        "entry.1342537331": values['id']
    }

def send_attendance(url, data):
    count = 0
    while count < 3:
        count += 1
        try:
          r = requests.post(url, data=data)
          break
        except:
          print("Error Occured!")
          time.sleep(2)

def validate(current_question_output, exercise_number, exercise_extra):
  student_email=!gcloud config get-value account
  if not student_email or 'unset' in student_email[0]:
    !gcloud auth login
    student_email=!gcloud config get-value account
  
  if current_question_output:
      exercise_score = True
      print("Parabéns")
  else:
      exercise_score = False
      print("Resposta incorreta.")

  values = {"exercise_number": exercise_number, "student_id": student_email[0],
            "exercise_points": 1, "exercise_score": exercise_score, 
            "exercise_extra": exercise_extra,
            "id": f"{student_email[0]}_{exercise_number}"}
  final_data = format_values(values)
  send_attendance(f"{url}&emailAddress={quote(str(student_email[0]))}", final_data)

```

### 2.3 Escreva uma função para contar a quantidade de strings, presentes numa lista, cujo tamanho é maior que 2 e o primeiro e último caracteres são os mesmos

#### Exemplo:  ['abc', 'xyz', 'aba', '1221']
#### Resultado esperado : 2


<!-- grade_id: cell-5420c3983940f212, points: 0, grade: false, solution: true -->
```python
# Dica 1: A função len retorna o tamanho da string. Exemplo: len('abc'), retorna 3
# Dica 2: Uma string é uma lista de caracteres
def contagem(lista):
    ### BEGIN SOLUTION
    n = 0
    for s in lista:
        if len(s) > 2 and s[0] == s[-1]:
            n += 1
    return n
    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-f86ec02e2bcb84c6, points: 1, grade: true, solution: false -->
```python
# Validação para nota
v1 = contagem(['abc', 'xyz', 'aba', '1221']) == 2
v2 = contagem(['ac', '1111', 'aasdfsdfdda', '12ewwww21']) == 3
v3 = contagem(['', 'xyzaaaaaaaaa', 'aa', 'babababa']) == 0
current_question_output = v1 and v2 and v3
assert current_question_output
validate(current_question_output, "2.3", True)
```

### 2.4 Crie um código python para imprimir o seguinte padrão. 
#### 1 
#### 1 2 
#### 1 2 3 
#### 1 2 3 4 
#### 1 2 3 4 5

<!-- grade_id: cell-61d03fd9d1b0eb4e, points: 1, grade: true, solution: true -->
```python
## Essa questão não tem assert.
## Faça aqui
### BEGIN SOLUTION
## Solução sem utilizar funções características do python (demanda menos conhecimento)
for i in range(1,6):
    line = ""
    for j in range(1, i+1):
        line = line + str(j) + " "
    print(line)
    
## Solução mais elegante utilizando conhecimentos aprimorados do python
# O asterisco usado antes do argumento de uma função transforma uma lista em vários argumentos.
# Exemplo: print(*[1, 2, 3]) é traduzido para print(1, 2, 3)
for i in range(1,6):
    print(*list(range(1,i+1)))
### END SOLUTION
```

### 2.5 Escreva uma função Python para percorrer uma lista e retornar os números multiplos de N. Pare a iteração no ponto em que encontrar um número maior do que K. Não retorne o 1o número maior do que K.
#### Exemplo:  lista = [120, 115, 12, 32, 56, 200, 715, 152, 130, 155, 183, 300], N=5, K=200
#### Resultado Esperado: [120, 115, 200]

<!-- grade_id: cell-1ff4485b8100a753, points: 0, grade: false, solution: true -->
```python
def filtra_lista(lista, N, K):
    """retorne uma lista com todos os números multiplos de N menores do que K, 
       até o primeiro número maior do que K, não inclua este número na lista de resultado.
    """
    ### BEGIN SOLUTION
    result = []
    for x in lista:
        if x % N == 0 and x <= K:
            result.append(x)
        if x > K:
            break
    return result
    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-b6c3a5db7e709b2a, points: 1, grade: true, solution: false -->
```python
# Validação para nota
v1 = filtra_lista([120, 115, 12, 32, 56, 200, 715, 152, 130, 155, 183, 300], 5, 200) == [120, 115, 200]
current_question_output = v1
assert current_question_output
validate(current_question_output, "2.5", True)
```

### 2.6 Dado um número inteiro qualquer, escreva uma função python para retornar a quantidade de caracteres (dígitos) contida neste número.
##### Exemplo:  qtd_digitos(2345653) deve retornar 7

<!-- grade_id: cell-e69a127f71184645, points: 0, grade: false, solution: true -->
```python
def qtd_digitos(N):
    """retorne a quantidade de dígitos do número N.
    """
    ### BEGIN SOLUTION
    return len(str(N))    
    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-418c099df61d368b, points: 1, grade: true, solution: false -->
```python
# Validação
v1 = qtd_digitos(2345653) == 7
current_question_output = v1
assert current_question_output
validate(current_question_output, "2.6", True)
```

### 2.7 Escreva uma função Python para somar os valores dos numeros presentes numa string de números separados por vírgula.
#### Exemplo:  "1.0,  4.0,  10.0,  8.0,  4.0,  20.0,  23.0"
#### Resultado esperado: 70

<!-- grade_id: cell-f6c50b3e165db993, points: 0, grade: false, solution: true -->
```python
def soma_string(texto):
    """retorne a soma dos numeros presentes na variavel texto
    """
    ### BEGIN SOLUTION
    lista = texto.split(",")
    soma = 0
    for i in lista:
        soma = soma + float(i)
    return soma
    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-401910203dadc4db, points: 1, grade: true, solution: false -->
```python
# Validação
var = "1.0,  4.0,  10.0,  8.0,  4.0,  20.0,  23.0"
v1 = soma_string(var) == 70.0
current_question_output = v1
assert current_question_output
validate(current_question_output, "2.7", True)
```

### 2.8 Crie um código python para imprimir o seguinte padrão. 
#### Dica: Use uma sintaxe de geração de números em ordem decrescente. Por exemplo ```range(5, 0, -1)``` gera um iterator para a sequência 5, 4, 3, 2, 1
#### 5 4 3 2 1 
#### 4 3 2 1 
#### 3 2 1 
#### 2 1 
#### 1

<!-- grade_id: cell-fc8267fa2d7e4c7c, points: 0, grade: false, solution: true -->
```python
### Este exercicio não tem assert
## Faça aqui

### BEGIN SOLUTION
## Solução sem utilizar funções características do python (demanda menos conhecimento)
for i in range(5,0, -1):
    line = ""
    for j in range(i, 0, -1):
        line = line + str(j) + " "
    print(line)
    
## Solução mais elegante utilizando conhecimentos aprimorados do python
# O asterisco usado antes do argumento de uma função transforma uma lista em vários argumentos.
# Exemplo: print(*[1, 2, 3]) é traduzido para print(1, 2, 3)
for i in range(5,0, -1):
    print(*list(range(i, 0, -1)))
### END SOLUTION
```

### 2.9 Usando um loop do tipo for, crie uma função python para inverter uma lista.
#### Exemplo:  inverter([120, 115, 12, 32, 56]) deve retornar [56, 32, 12, 115, 120]

<!-- grade_id: cell-513f2b547866e072, points: 0, grade: false, solution: true -->
```python
def inverter(lista):
    """retorne uma versão invertida da lista passada como argumento
    """
    ### BEGIN SOLUTION
    resultado = []
    for i in lista[::-1]:
        resultado.append(i)
    return resultado
    ### END SOLUTION
```

```python
### Este exercicio não tem assert
## Faça aqui

```

```python
v1 =  inverter([120, 115, 12, 32, 56]) == [56, 32, 12, 115, 120]
current_question_output = v1
v1 = current_question_output
validate(current_question_output, "2.9", True)
```

### 2.10 Escreva uma função python para substituir vírgula por ponto numa string e depois converta os números para float. Retorne a lista de números.
#### Dica: use a função replace: ```texto.replace("," , ".")``` Este código substitui as virgulas por pontos e retorna uma string com o resultado.
#### Exemplo:  "0,20 0,40 0,56 0,99"
#### Resultado esperado: [0.2, 0.4, 0.56, 0.99]

<!-- grade_id: cell-e0a4e5deea9187c4, points: 0, grade: false, solution: true -->
```python
def convert_to_float(texto):
    """retorne uma lista de numeros do tipo float
    """
    ### BEGIN SOLUTION
    alterado = texto.replace(",", ".").split(" ")
    resultado = []
    for i in alterado:
        resultado.append(float(i))
    return resultado
    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-30d325468b6f6a67, points: 1, grade: true, solution: false -->
```python
# Validação
text = "0,20 0,40 0,56 0,99"
v1 = convert_to_float(text) == [0.2, 0.4, 0.56, 0.99]
current_question_output = v1
assert current_question_output
validate(current_question_output, "2.10", True)
```

### 2.11 Escreva uma função em python para imprimir o padrão a seguir. A função deve aceitar como argumento o tamanho da maior string de asteriscos.
#### Dica: Qual o resultado de se multiplicar uma string por um número em python?

``` 
 * 
 * * 
 * * * 
 * * * * 
 * * * * * 
 * * * * 
 * * * 
 * * 
 *
```

<!-- grade_id: cell-cd8f3abd73a5fcc4, points: 0, grade: false, solution: true -->
```python
### Este exercício não tem assert
## Faça aqui

### BEGIN SOLUTION
## Solução 1: Com dois laços for
for i in range(1, 6):
    print("* "*i)

for i in range(4, 0, -1):
    print("* "*i)

## Solução 2: Com um laço apenas
for i in range(1, 10):
    if i <= 5:
        print("* "*i)
    else:
        print("* "*(10-i))
### END SOLUTION
```

### 2.12 Escreva uma função que receba como entrada uma string conforme apresentada a seguir e retorne uma lista contendo a população dos 3 Estados brasileiros mais populosos.

```"""Acre=881935; Alagoas=3337357; Amapá=845731; Amazonas=4144597; Ceará=9132078;
 Distrito Federal=3015268; Espírito Santo=4018650; Goiás=7018354; Maranhão=7075181;
 Mato Grosso=3484466; Mato Grosso do Sul=2778986; Minas Gerais=21168791;
 Paraíba=4018127; Paraná=11433957; Pernambuco=9557071; Piauí=3273227; Rio de Janeiro=17264943;
 Rio Grande do Norte=3506853; Rio Grande do Sul=11377239; Rondônia=1777225;
 Roraima=605761; Santa Catarina=7164788; São Paulo=45919049; Sergipe=2298696;
 Tocantins=1572866; Bahia=14873064; Pará=8602865"""```

#### Dica 1: Nesta string, a informação de cada Estado está separada por um ;
#### Dica 3: A função split quebra uma string numa lista baseado no caracter de separação especificado. 
####     Exemplo: ```'1,2,3'.split(',')``` retorna a lista ```['1','2','3']```
#### Dica 4: a função built-in ```.sort``` de uma lista ordena a referida lista ordenada 

<!-- grade_id: cell-9e2eb40f94d9f5e4, points: 0, grade: false, solution: true -->
```python
def top3_estados(lista_texto):
    """retorne os 3 estados de maior população. O parâmetro de entrada deve ser uma string
         conforme apresentado na descrição do exercício.
    """
    ### BEGIN SOLUTION
    lista1 = lista_texto.split(";")
    populacao = []
    for l in lista1:
        pop = l.split("=")[1]
        populacao.append(int(pop))
    populacao.sort()
    return populacao[-3:]
    ### END SOLUTION
```

```python
# Teste aqui

```

```python
# Validação
lista_estados="""Acre=881935; Alagoas=3337357; Amapá=845731; Amazonas=4144597; Ceará=9132078;\
 Distrito Federal=3015268; Espírito Santo=4018650; Goiás=7018354; Maranhão=7075181;\
 Mato Grosso=3484466; Mato Grosso do Sul=2778986; Minas Gerais=21168791;\
 Paraíba=4018127; Paraná=11433957; Pernambuco=9557071; Piauí=3273227; Rio de Janeiro=17264943;\
 Rio Grande do Norte=3506853; Rio Grande do Sul=11377239; Rondônia=1777225;\
 Roraima=605761; Santa Catarina=7164788; São Paulo=45919049; Sergipe=2298696;\
 Tocantins=1572866; Bahia=14873064; Pará=8602865"""

v1 = set(top3_estados(lista_estados)) == {45919049, 21168791, 17264943}
current_question_output = v1
assert current_question_output
validate(current_question_output, "2.12", True)
```

### 2.13 Escreva uma função Python que aceite como entrada uma string de palavras separadas por vírgula e retorne uma string com as palavras em ordem alfabética e separadas por vírgula. 
#### Exemplo:  Para t="hoje,nota,teste,bolsa"
#### Resultado Esperado: "bolsa,hoje,nota,teste"
#### Dica: use as funções built-in ```.split``` e ```.join``` presente nas strings


<!-- grade_id: cell-3201ecbb117ca859, points: 0, grade: false, solution: true -->
```python
def sort_words(words):
    """retorne uma string ordenada de palavras separadas por vírgula
    """
    ### BEGIN SOLUTION
    return ",".join(sorted(words.split(",")))
    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

```python
# Validação
v1 = sort_words("hoje,nota,teste,bolsa") == "bolsa,hoje,nota,teste"
current_question_output = v1
assert current_question_output
validate(current_question_output, "2.13", True)
```