# Aula 2- Exercícios - Sintaxe de Python e NumPy Array

### 2.1 Escreva uma função em Python que aceite um inteiro (n) de 1 a 9 e compute o valor de: n+nn+nnn
#### Valide os valores de entrada. Para isso, use a instrução raise ValueError
##### Exemplo: n=5. Então n + nn + nnn = 5 + 55 + 555 = 615

<!-- grade_id: cell-a2f686fdd79e4f51, points: 0, grade: false, solution: true -->
```python
def calculo_n(n):
    """retorne o valor de: n+nn+nnn
    """
    ### BEGIN SOLUTION
    return 123 * n
    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-5412543179b17057, points: 1, grade: true, solution: false -->
```python
# Validação
assert calculo_n(1) == 123
assert calculo_n(2) == 246
assert calculo_n(3) == 369
assert calculo_n(9) == 1107
```

### 2.2 Escreva uma função para retornar o maior item de uma lista de inteiros

<!-- grade_id: cell-2e044a5f0541da37, points: 0, grade: false, solution: true -->
```python
def maior(lista):
    """retorne o maior item da lista
    """
    ### BEGIN SOLUTION
    maior = lista[0]
    for i in lista[1:]:
        if i > maior:
            maior = i
    return maior
    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-5bd4079b37b6a7f3, points: 1, grade: true, solution: false -->
```python
# Validação
assert maior([1,0,-1,100,7]) == 100
assert maior([-100,-1580,-1,-9100,-77]) == -1
```