# Aula 2- Exercícios - Sintaxe de Python e NumPy Array

```python
#@title Init Cell - Execute-me!!!
import sys
from datetime import datetime
import requests
import time
from requests.utils import quote
url = 'https://docs.google.com/forms/d/e/1FAIpQLSeASSC8-w8FmfodZ4lBnuSEAvYuE4vatIBowLIREG1f-2pIpA/formResponse?usp=pp_url&entry.1986154915=mbacd&entry.513694412=2021&entry.1914621244=CienciaDeDados'

def format_values(values):
    return {
        "entry.1269959472": values['student_id'],
        "entry.1799867692": str(values['exercise_number']).replace(".", "_"),
        "entry.886231469": values['exercise_score'],
        "entry.1491599254": values['exercise_extra'],
        "entry.1342537331": values['id']
    }

def send_attendance(url, data):
    count = 0
    while count < 3:
        count += 1
        try:
          r = requests.post(url, data=data)
          break
        except:
          print("Error Occured!")
          time.sleep(2)

def validate(current_question_output, exercise_number, exercise_extra):
  student_email=!gcloud config get-value account
  if not student_email or 'unset' in student_email[0]:
    !gcloud auth login
    student_email=!gcloud config get-value account
  
  if current_question_output:
      exercise_score = True
      print("Parabéns")
  else:
      exercise_score = False
      print("Resposta incorreta.")

  values = {"exercise_number": exercise_number, "student_id": student_email[0],
            "exercise_points": 1, "exercise_score": exercise_score, 
            "exercise_extra": exercise_extra,
            "id": f"{student_email[0]}_{exercise_number}"}
  final_data = format_values(values)
  send_attendance(f"{url}&emailAddress={quote(str(student_email[0]))}", final_data)

```

### Ao resolver um exercício, apague o seguinte trecho do código:
```# YOUR CODE HERE
raise NotImplementedError()
```

### 2.1 Escreva uma função em Python que aceite um inteiro (n) de 1 a 9 e compute o valor de: n+nn+nnn
#### Valide os valores de entrada. Para isso, use a instrução raise ValueError
##### Exemplo: n=5. Então n + nn + nnn = 5 + 55 + 555 = 615

<!-- grade_id: cell-a2f686fdd79e4f51, points: 0, grade: false, solution: true -->
```python
def calculo_n(n):
    """retorne o valor de: n+nn+nnn
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-5412543179b17057, points: 1, grade: true, solution: false -->
```python
# Validação para nota
v1 = calculo_n(1) == 123
v2 = calculo_n(2) == 246
v3 = calculo_n(3) == 369
v4 = calculo_n(9) == 1107
current_question_output = v1 and v2 and v3 and v3
assert current_question_output
validate(current_question_output, "2.1", False)
```

### 2.2 Escreva uma função para retornar o maior item de uma lista de inteiros

<!-- grade_id: cell-2e044a5f0541da37, points: 0, grade: false, solution: true -->
```python
def maior(lista):
    """retorne o maior item da lista
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-5bd4079b37b6a7f3, points: 1, grade: true, solution: false -->
```python
# Validação para nota
v1 = maior([1,0,-1,100,7]) == 100
v2 = maior([-100,-1580,-1,-9100,-77]) == -1
current_question_output = v1 and v2
assert current_question_output
validate(current_question_output, "2.2", False)
```