# Aula 3 - Exercícios Extra - Estruturas de Dados Nativas do Python


```python
#@title Init Cell - Execute-me!!!
import sys
from datetime import datetime
import requests
import time
from requests.utils import quote
url = 'https://docs.google.com/forms/d/e/1FAIpQLSeASSC8-w8FmfodZ4lBnuSEAvYuE4vatIBowLIREG1f-2pIpA/formResponse?usp=pp_url&entry.1986154915=mbacd&entry.513694412=2021&entry.1914621244=CienciaDeDados'

def format_values(values):
    return {
        "entry.1269959472": values['student_id'],
        "entry.1799867692": str(values['exercise_number']).replace(".", "_"),
        "entry.886231469": values['exercise_score'],
        "entry.1491599254": values['exercise_extra'],
        "entry.1342537331": values['id']
    }

def send_attendance(url, data):
    count = 0
    while count < 3:
        count += 1
        try:
          r = requests.post(url, data=data)
          break
        except:
          print("Error Occured!")
          time.sleep(2)

def validate(current_question_output, exercise_number, exercise_extra):
  student_email=!gcloud config get-value account
  if not student_email or 'unset' in student_email[0]:
    !gcloud auth login
    student_email=!gcloud config get-value account
  
  if current_question_output:
      exercise_score = True
      print("Parabéns")
  else:
      exercise_score = False
      print("Resposta incorreta.")

  values = {"exercise_number": exercise_number, "student_id": student_email[0],
            "exercise_points": 1, "exercise_score": exercise_score, 
            "exercise_extra": exercise_extra,
            "id": f"{student_email[0]}_{exercise_number}"}
  final_data = format_values(values)
  send_attendance(f"{url}&emailAddress={quote(str(student_email[0]))}", final_data)

```

### 3.3 Escreva uma função Python para combinar os itens de mesma chave de uma lista de  dicionários, somando os valores correspondentes.
#### Exemplo:  para ```[{'item': 'carro', 'quantidade': 400}, {'item': 'moto', 'quantidade': 300}, {'item': 'carro', 'quantidade': 750}]```
#### Resultado esperado: ```{'carro': 1150, 'moto': 300}```

<!-- grade_id: cell-e0a4e5deea9187c4, points: 0, grade: false, solution: true -->
```python
def combine_items(item_list):
    """retorne um dicionário combinando (somando) as quantidades do mesmo item
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-30d325468b6f6a67, points: 1, grade: true, solution: false -->
```python
# Validação
input_data = [{'item': 'carro', 'quantidade': 400}, 
              {'item': 'moto', 'quantidade': 300}, 
              {'item': 'carro', 'quantidade': 750}]
output_result = {'carro': 1150, 'moto': 300}
v1 = combine_items(input_data) == output_result
current_question_output = v1
assert current_question_output
validate(current_question_output, "3.3", True)
```

### 3.4 Usando um defaultdict, escreva uma função Python para transformar uma lista de listas num dicionário. Os valores repetidos de uma lista devem ser associados a mesma chave do dicionário, como no exemplo abaixo.
#### Exemplo:  ```[ ["carro", 10], ["moto", 20], ["trem", 5], ["carro", 30], ["caminhao", 20] ]```
#### Resultado esperado: ```{"carro": [10, 30], "moto": [20], "trem": [5], "caminhao": [20] }```

<!-- grade_id: cell-9959d8353ef5e348, points: 0, grade: false, solution: true -->
```python
def dict_from_list(lista):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
## Faça seus testes aqui

```

```python
# Validação
lista2 = [ ["carro", 10], ["moto", 20], ["trem", 5], ["carro", 30], ["caminhao", 20] ]
v1 = dict_from_list(lista2) == {'carro': [10, 30], 'moto': [20], 'trem': [5], 'caminhao': [20]}
current_question_output = v1
assert current_question_output
validate(current_question_output, "3.4", True)
```

### 3.5 Escreva uma função Python para computar o valor total de uma compra de supermercado e dar baixa no estoque (deduzir a quantidade comprada da quantidade em estoque). Retorne o valor total da compra e altere o dicionário de estoque sem retorná-lo.
#### Se a quantidade do item em estoque for menor do que a quantidade da lista de compra, o estoque restante deve ser vendido.

#### Exemplo:  
####                   ```lista_compras = {"arroz":1, "feijao":2, "macarrao":3}```
####                   ```estoque = {"arroz": 60, "feijao": 30, "macarrao": 1, "farinha": 15}```
####                   ```precos = {"arroz": 4, "feijao": 2, "macarrao": 1.5, "farinha": 3}```
#### Resultado esperado: ```total = 9.5```

<!-- grade_id: cell-fb745f3563064d90, points: 0, grade: false, solution: true -->
```python
def compra(estoque, precos, lista):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
## Faça seus testes aqui

```

```python
# Validação
lista_compras = {"arroz":1, "feijao":2, "macarrao":3}
estoque = {"arroz": 60, "feijao": 30, "macarrao": 1, "farinha": 15}
precos = {"arroz": 4, "feijao": 2, "macarrao": 1.5, "farinha": 3}
v1 = compra(estoque, precos, lista_compras) == 9.5
current_question_output = v1
assert current_question_output
validate(current_question_output, "3.5", True)
```

### 3.6 Usando o tipo de dados set (conjuntos) do python, crie uma função python que  retorne os elementos que não estão na intersecção das duas listas passadas como argumento. O retorno deve ser uma lista.
#### Exemplo:
```
lista1 = ["awerr", "aa", "awyt", "g", "sssaere"]
lista2 = ["awerr", "aa", "sssaere", "aabb", "qwerttt", "12345"]```
#### Retorno esperado:
```
['awyt', 'g', 'aabb', 'qwerttt', '12345']```
#### dica 1: Converta as listas para set. Exemplo: ```c1=set(lista1)```
#### dica 2: crie o conjunto interseção e o conjunto uniao dos dois conjuntos correspondentes às duas listas.
#### dica 3: faça a diferença entre a união e a interseção.
#### dica 4: faça a conversão do conjunto resultado para lista: ```resultado = list(c)```



<!-- grade_id: cell-06bc42e7cd22bdcf, points: 0, grade: false, solution: true -->
```python
def not_union_inters(l1, l2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
## Faça seus testes aqui

```

```python
lista1 = ["awerr", "aa", "awyt", "g", "sssaere"]
lista2 = ["awerr", "aa", "sssaere", "aabb", "qwerttt", "12345"]
v1 = set(not_union_inters(lista1, lista2)) == {'12345', 'qwerttt', 'awyt', 'aabb', 'g'}
current_question_output = v1
assert current_question_output
validate(current_question_output, "3.6", True)
```

### 3.7 Crie uma função para gerar tuplas compostas pelos elementos em posições de mesmo ínidce numa lista e na sua lista reversa. A lista deve ser passada como argumento da função.
#### Exemplo: 
```lista = [1, 4, 6, 10, 5]
lista_reversa = [5, 10, 6, 4, 1]```
#### Resultado esperado: 
```[(1,5) , (4,10) , (6,6) , (10,4) , (5,1)]
```
#### Dica: use a função zip

<!-- grade_id: cell-c651c0d66cc0aeb3, points: 0, grade: false, solution: true -->
```python
def lista_tuplas(l):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
## Faça seus testes aqui

```

```python
lista = [1, 4, 6, 10, 5]
v1 = lista_tuplas(lista) == [(1, 5), (4, 10), (6, 6), (10, 4), (5, 1)]
current_question_output = v1
assert current_question_output
validate(current_question_output, "3.7", True)
```

### 3.8 Crie uma função python que receba uma lista como argumento, e usando um list comprehenion, crie uma lista de dicionários. A chave de cada dicionário, na lista resultante, deve ser o elemento da lista (passada como argumento) e o valor a sua posição (índice) dentro desta lista.
#### Exemplo: 
```lista = ["azul", "amarelo", "vermelho"]```
#### Resultado esperado: 
```[{"azul": 0}, {"amarelo": 1}, {"vermelho": 2}]```

#### Dica: use a função index para extrair o índice de um elemento numa lista. Exemplo: ``` lista.index("azul") ```


<!-- grade_id: cell-07db5d73a369b436, points: 0, grade: false, solution: true -->
```python
def lista_dicionarios(l):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
## Faça seus testes aqui

```

```python
lista = ["azul", "amarelo", "vermelho"]
v1 = lista_dicionarios(lista) == [{'azul': 0}, {'amarelo': 1}, {'vermelho': 2}]
current_question_output = v1
assert current_question_output
validate(current_question_output, "3.8", True)
```

### 3.9 Crie uma função python que receba um dicionário (de alunos e notas) como argumento, e usando um list comprehenion, crie uma lista com as respectivas notas abaixo de um certo valor passado como argumento da função.
#### Exemplo: 
```lista = {"Joao": 10, "Jose": 5, "Roberta": 9, "Maria": 6}
minimo = 7```
#### Resultado esperado: 
```[5, 6]```

#### Dica: use a seguinte sintaxe de list comprehension ```[i for i in lista if condicao] ``` onde ```condição``` é uma expressão que retorna um valor ```True``` ou ```False```. Neste caso, estarão na lista resultante somente os elementos cuja ```condição``` retornou ```True```.


<!-- grade_id: cell-4afc26f32762755c, points: 0, grade: false, solution: true -->
```python
def nota_baixa(n, minimo):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
## Faça seus testes aqui

```

```python
notas = {"Joao": 10, "Jose": 5, "Roberta": 9, "Maria": 6}
minimo = 7
v1 = nota_baixa(notas, minimo) == [5, 6]
current_question_output = v1
assert current_question_output
validate(current_question_output, "3.9", True)
```

### Os exercícios a seguir se baseiam num pequeno extrato de um dataset real dos serviços públicos brasileiros. Os serviços a seguir são da Fio Cruz. 

#### Curiosidade 1: Se tiver curiosidade consultar a fonte diretamente, execute o seguinte comando num console: ```curl https://www.servicos.gov.br/api/v1/servicos/orgao/315```
#### Curiosidade 2: Se quiser visualizar esses resultados no navegador no formato XML, acesse esta URL no navegador https://www.servicos.gov.br/api/v1/servicos/orgao/315
#### Curiosidade 3: Se quiser visualizar esses resultados no navegador no formato JSON, acesse a URL https://jsonformatter.org/json-viewer, e use a função Load Data para carregar os dados da URL https://www.servicos.gov.br/api/v1/servicos/orgao/315.
#### Curiosidade 4: 315 é o código Siorg da Fiocruz
#### Curiosidade 5: A documentação desta API encontra-se publicada em https://www.servicos.gov.br/api/v1/docs

```python
# Execute esta célula para criar a variavel servicos que vai conter a lista de serviços
# Na disciplina D20 será aprofundado esse assunto de requisições HTTP e APIs com python e curl
# Importar a biblioteca que faz requisições HTTP
import requests
# Fazer uma requisição HTTP
r = requests.get("https://www.servicos.gov.br/api/v1/servicos/orgao/315")
# Transformar a resposta em JSON e pegar a lista de serviços cuja chave é "resposta"
servicos = r.json()["resposta"]
# Imprimir o primeiro serviço da lista
print(servicos[0])
```

### 3.10 Percorra a lista ```servicos``` selecionando os atributos de interesse ```nome, sigla, servicoDigital, nomeOrgao```. Construa o atributo ```id_servico``` a partir do atributo ```id``` coletando o número (id do serviço) que aparece depois da ultima barra da URL do serviço. Construa um dicionário de dicionários. No dicionário externo as chaves são o ```id_servico```, e seu respectivo valor deve ser um dicionário com os atributos de interesse.

<!-- grade_id: cell-787017167e839bb1, points: 0, grade: false, solution: true -->
```python
def atributos_interesse(serv):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
## Faça seus testes aqui

```

```python
atr_inter = atributos_interesse(servicos)
gabarito = {'nome': 'Solicitar diagnóstico de referência em chikungunya',
  'sigla': 'LABFLA', 'servicoDigital': False, 'nomeOrgao': 'Fundação Oswaldo Cruz (FIOCRUZ)'}
v1 = atr_inter['9031'] == gabarito
current_question_output = v1
assert current_question_output
validate(current_question_output, "3.10", True)
```

### 3.11 Crie uma função que receba como argumento o dicionário do exercício anterior e calcule a porcentagem de serviços digitais presentes no dicionário.

<!-- grade_id: cell-44e66bb6d2a495c1, points: 0, grade: false, solution: true -->
```python
def calc_porcentagem_digital(servs):
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

```python
## Faça seus testes aqui

```

```python
atr_inter = atributos_interesse(servicos)
v1 = round(calc_porcentagem_digital(atr_inter), 2) == 0.23
current_question_output = v1
assert current_question_output
validate(current_question_output, "3.11", True)
```